Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Static



By now you're familiar with the typical use of linear regression: you're given a dataframe with a bunch of features and a set of observations for each feature. You then take one feature (your y) and try to predict it based on other features (your x's). 

This generally works to a certain extent. This morning, though, we're going to set up a rather strange problem for linear regression. We will see what happens when y and the x's are totally unrelated.

Recall that regression follows the formula 
$$
y = \beta_0 + \beta_1 x_1 + \beta_2 x_2 + \dots + \epsilon
$$ 
When we "fit" an OLS model we are giving the model all of the x's and the y and asking it to find the best betas.

This time, however, we're going to set all the betas to zero 
$$
\beta_0=\beta_1=\dots=0
$$
and then fit an OLS model.

Before we do this, stop and think for a second:

- What do you expect the model to do? What will the betas/r-squared/p-values that it finds look like?
- What do you think the model *should* do? Is that different from what you think it *will* do?

![](https://upload.wikimedia.org/wikipedia/commons/5/5a/No_Signal_23.JPG)

# Part 1

Generate simulation data. We want to have 200 points (observations) for the y feature and for 20 x features. In other words make sure `y.shape == (200,1)` and `x.shape == (200,20)`. The x's should be randomly generated independent of each other. And the y should be randomly generated independent of the x's. 

Use statsmodels to fit an OLS model to your data. Are the results as you expected? Do you have any betas with a $p<0.05$? If not, re-run the model until you do.

```python
import statsmodels.api as sm
import statsmodels.formula.api as smf
import seaborn
seaborn.set_style("darkgrid")
import pandas as pd
%pylab inline
```

<!-- grade_id: cell-42f21cf570e8a3ba, points: 0, grade: false, solution: true -->
```python
n_points = 200
n_features = 20

# store the output of fit=model.fit()

# YOUR CODE HERE
raise NotImplementedError()

fit = model.fit()
fit.summary2()
```

<!-- grade_id: cell-374b565f4b042532, points: 10, grade: true, solution: false -->
```python
# correct shape for randomly generated data
assert y.shape == (200,1) 
assert x.shape == (200,20)
```

<!-- grade_id: cell-38b21ec7d724b529, points: 10, grade: true, solution: false -->
```python
# data is random
assert sum(fit.pvalues < .05) < 5
assert fit.rsquared < .3
assert fit.rsquared_adj < .1
```

```python

```

# Part 2

Now, automate the process! Run the above analysis but vary the number of x's from 1 to 200. Log the number of significant features, r2 and r2-adj for each case and plot them. Use $p<0.05$ for the cutoff for p values.

So, you will run the full analysis for `n_features=1`, then log the number of significant features, r2 and r2-adj. And then you will repeat that for `n_features=2`, `n_features=3`, and so on. When you are finished, you should have three lists, one each for the number of significant features, r2 and r2-adj. Each list should have 200 items in it.

Hint: You will need to get the r2, r2-adj, and the number of significant features from the fit object. There are two ways to find out more about what is stored in that object:
- running `dir(fit)` returns all of the attributes that can be accessed, such as `fit.mse_total` for the total Mean-Squared Erro. 
- The fit object is explained in more detail here in the [statsmodels documentation](https://www.statsmodels.org/stable/generated/statsmodels.regression.linear_model.RegressionResults.html#statsmodels.regression.linear_model.RegressionResults)

<!-- grade_id: cell-80550bb841262b8c, points: 0, grade: false, solution: true -->
```python
n_points = 200
n_features = 200
num_sig = [] # store the number of significant features in this list
rsquared = [] # store rsquared in this list
rsquared_adj = [] # store rsquared-adj in this list

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fc03df3bcc5d39fd, points: 10, grade: true, solution: false -->
```python
# all results are 200 long
assert all([len(l) == 200 for l in [num_sig, rsquared, rsquared_adj]])
```

<!-- grade_id: cell-2ce0a7ba49c3b502, points: 10, grade: true, solution: false -->
```python
assert all([len(l) == 200 for l in [num_sig, rsquared, rsquared_adj]])
assert mean(rsquared[-10:]) > .9
assert mean(rsquared_adj[:-1]) < .1
assert mean(num_sig) < 10
```

```python
fig,axes = subplots(nrows=3, ncols=1, sharex=True, figsize=(10,15))

axes[0].plot(num_sig)
axes[0].set_title("number of significant features")
axes[1].plot(rsquared)
axes[1].set_title("$R^2$")
axes[2].plot(rsquared_adj)
axes[2].set_title("$R^2_{adj}$");
```

Finally in the cell below, in a few sentences each, answer the following questions:
- How did the model's performance change as you added more features?
- Since you generated the data, you know the correct answer for the number of significant features. What is it?
- By the end, it looks like there is a large difference between the actual performance of the model and the rsquared measure of performance. If you needed to report an r-squared value and r-squared adjusted wasn't acceptable, how would you find an accurate way to measure the model's performance?

<!-- grade_id: cell-c5dd3e7a3f56c5fa, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE